#### Exercise 2 (Lists)
Find the sum of the smallest 10 numbers of List L Indexing, Sorted and sum function could be useful

<!-- grade_id: cell-d22595cf41ad5ff3, points: 0, grade: false, solution: true -->
```python
L =[75, 28, 47, 51, 96, 76, 32, 57, 10, 10, 25, 62, 17,  9, 68, 65, -2,
       54, 34, 41, 74, 83, 91, 59, -3, 88, 59, 91, 10, -9, 90, 46, 58, 25,
       62, 32, 74,  1, 61, 43, 25, 62, -5, 49, 40, 44, 83, -5, 13, 35]

### BEGIN SOLUTION
ans = sum(sorted(L)[:10])
### END SOLUTION

```

<!-- grade_id: cell-c9417b8946afbe0c, points: 10, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ####
assert(ans==16)
### BEGIN HIDDEN TESTS
assert(ans==sum(sorted(L)[:10]) )
### END HIDDEN TESTS
```

```python

```