#### Assignment 1.1 Python
Python is an easy to learn, powerful programming language with efficient high-level data structures and object-oriented programming. Python’s elegant syntax and dynamic typing, together with its interpreted nature, make it an ideal language for scripting and rapid application development in many areas on most platforms.

Exercise 1
T is a Nested Tuple, use indexing to extract 'g' from the Tuple

#### Exercise 1
T is a Nested Tuple, use indexing to extract 'g' from the Tuple

<!-- grade_id: cell-6aab945b0c90a31a, points: 0, grade: false, solution: true -->
```python
T = ('a','b','c',('d','e'),('f',('f','g'),'k'),'p')
thisisg = None
### BEGIN SOLUTION
thisisg = T[4][1][1]
### END SOLUTION
```

<!-- grade_id: cell-3f69882e01a929b5, points: 10, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ####
assert(thisisg=='g')
### BEGIN HIDDEN TESTS
assert(thisisg=='g')
### END HIDDEN TESTS
```